<h1>A-Star Search Solver for the 8-puzzle Game</h1>

---

<h3>
The Puzzle Node   
</h3>

<!-- grade_id: b5da05, points: 0, grade: false, solution: true -->
```python
# Adapted from 
# A* Tree Search Example for Robot Navigation
# R. Shekhar
# August 10, 2017

# import necessary libraries
import copy
from copy import deepcopy

class PuzzleNode:
    """
     A class representing a node in a puzzle game.
    """
    def __init__(self,state,fval,gval,parent=None):
        """
        The constructor method
        state: the state of the node
        fval: the f-value of the node
        gval: the g-value of the node
        parent: the parent of the node
        """
        self.state = state
        self.fval = fval
        self.gval = gval
        self.parent = parent

    def __lt__(self,other):
        """
        Method that compares the "fval" attribute of two objects of the same class.
        Returns True if the "fval" attribute of the current object is less than 
        the other object's "fval" attribute. 
        """
        if self.fval < other.fval:
            return True
        else:
            return False

    def __str__(self):
        """
        Method that iterates over the rows of the state.
        Returns a string representation of the state.
        """
        # initialize an empty string
        state_string = ""
        # iterate over each row of the state and add it to state_string
        for row in self.state:
            state_string += str(row) + "\n"
        return state_string
    
    def dimension_isValid(self):
        """
        Method that iterates over the rows of the state and checks the dimensions.
        Returns True if dimensions of the state are valid. 
        """
        # get dimension of the state
        n = len(self.state)
        # create set of valid entries in the state
        numbers = set(range(n*n))
        # represent the state as a list
        states_list = [number for row in self.state for number in row]
        # check if the length of the list is equal to n*n 
        # check if the set of entries in the state list is equal to the set of valid numbers
        if len(states_list) != n*n or set(states_list) != numbers:
            return False
        else:
            return True
        
    def puzzle_isSolvable(self):
        """
        Method that checks whether the start state configuration is solvable.
        Returns True if the puzzle is solvable.
        """
        # represent the state as a list
        states_list = [number for row in self.state for number in row]
    
        # count the number of out of order tiles
        out_of_order = 0
        for i in range(len(states_list)):
            for j in range(i + 1, len(states_list)):
                # check whether both tiles exist
                # checks whether the tile at index i is larger than the tile at index j
                if states_list[i] and states_list[j] and states_list[i] > states_list[j]:
                    out_of_order += 1
        # check whether the puzzle is solvable
        if len(states_list)**0.5 % 2 == 1:
            # if the puzzle has an odd number of rows
            # check whether the number of out of order tiles is even
            return out_of_order % 2 == 0
        else:
            # if the puzzle has an even number of rows
            # check whether the blank tile is on an odd or even row
            blank_tile_position = len(states_list)**0.5 - (states_list.index(0) // len(states_list)**0.5)
            if blank_tile_position % 2 == 0:
                return False
            else:
                return True

        
    def zero_tile(self):
        """
        Method that searches the state for the position of the blank tile.
        Returns the position as a tuple of row and column indices.
        """
        for i in range(len(self.state)):
            for j in range(len(self.state)):
                # check whether the current tile is the blank tile
                if self.state[i][j] == 0:
                    return i, j
                
    def create_goal_state(self):
        """
        Method that creates a goal state given a start state configuration.
        Returns a matrix representation of the goal state
        """
        # initialize an empty list to store the goal state
        goal_state = []
        # initialize a count to track the cell values
        count = 0
        for i in range(len(self.state)):
            # initialize an empty list to store each row of the goal state
            row = []
            for j in range(len(self.state)):
                # add the current count to the row
                row.append(count)
                count += 1
            # add the complete row to the goal state
            goal_state.append(row)
        # set the top-left corner to 0
        goal_state[0][0] = 0
        return goal_state
    

    def get_successors(self):
        """
        Method that creates the possible next moves of the current state.
        Returns list of lists representations of possible successor states.
        """
        # initialize empty list to store successors
        successors = []
        # size of the start state
        n = len(self.state)
        # get the row and column of the zero tile
        zero_row, zero_col = self.zero_tile()
        
        # if the zero tile can move up
        if zero_row > 0:
            # create a deep copy of the current state
            new_state = copy.deepcopy(self.state)
            # swap zero with the tile above it
            new_state[zero_row][zero_col] = new_state[zero_row - 1][zero_col]
            new_state[zero_row - 1][zero_col] = 0
            # add the new state of the list of successors
            successors.append(PuzzleNode(new_state, self.fval + 1, self.gval + 1, self))
        
        # if the zero tile can move down
        if zero_row < n - 1:
            # create a deep copy of the current state
            new_state = copy.deepcopy(self.state)
            # swap zero with the tile below it
            new_state[zero_row][zero_col] = new_state[zero_row + 1][zero_col]
            new_state[zero_row + 1][zero_col] = 0
            # add the new state of the list of successors
            successors.append(PuzzleNode(new_state, self.fval + 1, self.gval + 1, self))
        
        # if the zero tile can move left
        if zero_col > 0:
            # create a deep copy of the current state
            new_state = copy.deepcopy(self.state)
            # swap zero with the tile to its left
            new_state[zero_row][zero_col] = new_state[zero_row][zero_col - 1]
            new_state[zero_row][zero_col - 1] = 0
            # add the new state of the list of successors
            successors.append(PuzzleNode(new_state, self.fval + 1, self.gval + 1, self))
        
        # if the zero tile can move right
        if zero_col < n - 1:
            # create a deep copy of the current state
            new_state = copy.deepcopy(self.state)
            # swap zero with the tile to its right
            new_state[zero_row][zero_col] = new_state[zero_row][zero_col + 1]
            new_state[zero_row][zero_col + 1] = 0
            # add the new state of the list of successors
            successors.append(PuzzleNode(new_state, self.fval + 1, self.gval + 1, self))
        return successors

#raise NotImplementedError()
```

<h3>
The Heuristics  
</h3>

<!-- grade_id: 3fb850, points: 0, grade: false, solution: true -->
```python
def h1(state):
    """
    Function that compares the start state and goal state, and checks for misplaced tiles.
    Returns the number of misplaced tiles as the heuristic value.
    """
    # initialize heuristic value to zero
    h = 0
    # create a start state
    node = PuzzleNode(state, 0, 0)
    # create a goal state
    goal_state = node.create_goal_state()
    
    for i in range(len(state)):
        for j in range(len(state)):
            # check if the cell entry is not zero and is in the wrong position
            if state[i][j] != 0 and state[i][j] != goal_state[i][j]:
                # increase the heuristic value by 1
                h += 1
    return h

def h2(state):
    """
    Function that compares the start state and goal state and calculates the distance between them.
    Returns the Manhattan distance from the solved configuration
    """

    # initialize the heuristic value to zero
    h = 0
    
    for i in range(len(state)):
        for j in range(len(state)):
            # check if the cell entry is not zero
            if state[i][j] != 0:
                # calculate the row where the entry should be
                row = state[i][j] % len(state)
                # calculate the col where the entry should be
                col = state[i][j] // len(state)
                # add the distance of the entry from its current position
                # to its correct position to the manhattan distance
                h += abs(row - j) + abs(col - i)
    return h

#raise NotImplementedError()
    
# Extra heuristic for the extension.  If implemented, modify the function below
def h3(state):
    """
    This function returns a heuristic that dominates the Manhattan distance, given the board state
    Input:
        -state: the board state as a list of lists
    Output:
        -h: the Heuristic distance of the state from its solved configuration
    """
    return 0

# If you implement more than 3 heuristics, then add any extra heuristic functions onto the end of this list.
heuristics = [h1, h2, h3]


```

<h3>
A-Star Search Algorithm  
</h3>

<!-- grade_id: 328d7b, points: 0, grade: false, solution: true -->
```python
import heapq

def solvePuzzle(state, heuristic):
    """This function should solve the n**2-1 puzzle for any n > 2 (although it may take too long for n > 4)).
    Inputs:
        -state: The initial state of the puzzle as a list of lists
        -heuristic: a handle to a heuristic function.  Will be one of those defined in Question 2.
    Outputs:
        -steps: The number of steps to optimally solve the puzzle (excluding the initial state)
        -exp: The number of nodes expanded to reach the solution
        -max_frontier: The maximum size of the frontier over the whole search
        -opt_path: The optimal path as a list of list of lists.  That is, opt_path[:,:,i] should give a list of lists
                    that represents the state of the board at the ith step of the solution.
        -err: An error code.  If state is not of the appropriate size and dimension, return -1.  For the extention task,
          if the state is not solvable, then return -2
    """
    
    # determine the state dimension n
    n = len(state)
    
    # initialize output variables
    steps = 0
    exp = 0
    max_frontier = 1
    opt_path = []

    # initialize start node and goal state
    start_node = PuzzleNode(state, heuristic(state), 0)
    goal_state = [[i*n+j for j in range(n)] for i in range(n)]
    
    # check if the state has correct dimensions
    if start_node.dimension_isValid() == False:
        steps, exp, max_frontier, opt_path, err = 0, 0, 1, [], -1
        return steps, exp, max_frontier, opt_path, err
    
    # check if the state is solvable
    if start_node.puzzle_isSolvable() == False:
        steps, exp, max_frontier, opt_path, err = 0, 0, 1, [], -2
        return steps, exp, max_frontier, opt_path, err

    # initialize frontier and explored set
    frontier = [start_node]
    # does not return the already explored paths
    explored = set()


    # A* search algorithm
    while frontier:
        # update maximum frontier size
        max_frontier = max(max_frontier, len(frontier))

        # choose node with lowest f-value
        current_node = heapq.heappop(frontier)

        # check if goal state has been reached
        if current_node.state == goal_state:
            break

        # add current node to explored set
        explored.add(str(current_node))

        # expand current node
        # get the successor nodes
        successors = current_node.get_successors()
        for successor_node in successors:
            if str(successor_node.state) not in explored:
                # g-value for the successor node
                successor_gval = current_node.gval + 1
                # f-value for the successor node
                successor_fval = successor_gval + heuristic(successor_node.state)
                # assign attributes
                successor_node.gval = successor_gval
                successor_node.fval = successor_fval
                successor_node.parent = current_node
                # add successor node to priority queue
                heapq.heappush(frontier, successor_node)
                explored.add(str(successor_node.state))
            
        exp += 1
    
        
    # reconstruct optimal path
    # traverse the node's parents to get the path from start to goal
    while current_node:
        opt_path.append(current_node.state)
        current_node = current_node.parent
    # reversed to get the correct order from start to goal
    opt_path.reverse()
    # subtract 1 because the path is 1 less than the total states
    steps = len(opt_path) - 1
        
    if not opt_path:
        # return error code -2 if goal state is not found
        return steps, exp, max_frontier, opt_path, -2
    else:
        return steps, exp, max_frontier, opt_path, 0
       
#raise NotImplementedError()

```

```python
# SOLVE THE PUZZLE USING HEURISTIC 1

# define the start and goal states of the 8-puzzle game
start_state = [[1, 2, 3],
               [4, 0, 6],
               [7, 5, 8]]

# start state as a PuzzleNode instance
node = PuzzleNode(start_state, 0, 0)
# goal states as a PuzzleNode instance
goal_state = node.create_goal_state()
# solve the puzzle
steps, exp, max_frontier, opt_path, err = solvePuzzle(start_state, h1)

print("Steps:", steps)
print("Nodes expanded:", exp)
print("Max frontier size:", max_frontier)
print("Optimal path:", opt_path)
print("Error code:", err)
```

Output:
```
Steps: 20
Nodes expanded: 4131
Max frontier size: 2444
Optimal path: [[[1, 2, 3], [4, 0, 6], [7, 5, 8]], [[1, 2, 3], [4, 6, 0], [7, 5, 8]], [[1, 2, 0], [4, 6, 3], [7, 5, 8]], [[1, 0, 2], [4, 6, 3], [7, 5, 8]], [[0, 1, 2], [4, 6, 3], [7, 5, 8]], [[4, 1, 2], [0, 6, 3], [7, 5, 8]], [[4, 1, 2], [6, 0, 3], [7, 5, 8]], [[4, 1, 2], [6, 3, 0], [7, 5, 8]], [[4, 1, 0], [6, 3, 2], [7, 5, 8]], [[4, 0, 1], [6, 3, 2], [7, 5, 8]], [[4, 3, 1], [6, 0, 2], [7, 5, 8]], [[4, 3, 1], [6, 5, 2], [7, 0, 8]], [[4, 3, 1], [6, 5, 2], [0, 7, 8]], [[4, 3, 1], [0, 5, 2], [6, 7, 8]], [[0, 3, 1], [4, 5, 2], [6, 7, 8]], [[3, 0, 1], [4, 5, 2], [6, 7, 8]], [[3, 1, 0], [4, 5, 2], [6, 7, 8]], [[3, 1, 2], [4, 5, 0], [6, 7, 8]], [[3, 1, 2], [4, 0, 5], [6, 7, 8]], [[3, 1, 2], [0, 4, 5], [6, 7, 8]], [[0, 1, 2], [3, 4, 5], [6, 7, 8]]]
Error code: 0
```

```python
# SOLVE THE PUZZLE USING HEURISTIC 2

# define the start and goal states of the 8-puzzle game
start_state = [[1, 2, 3],
               [4, 0, 6],
               [7, 5, 8]]

# start state as a PuzzleNode instance
node = PuzzleNode(start_state, 0, 0)
# goal states as a PuzzleNode instance
goal_state = node.create_goal_state()
# solve the puzzle
steps, exp, max_frontier, opt_path, err = solvePuzzle(start_state, h2)

print("Steps:", steps)
print("Nodes expanded:", exp)
print("Max frontier size:", max_frontier)
print("Optimal path:", opt_path)
print("Error code:", err)
```

Output:
```
Steps: 20
Nodes expanded: 389
Max frontier size: 251
Optimal path: [[[1, 2, 3], [4, 0, 6], [7, 5, 8]], [[1, 2, 3], [4, 6, 0], [7, 5, 8]], [[1, 2, 0], [4, 6, 3], [7, 5, 8]], [[1, 0, 2], [4, 6, 3], [7, 5, 8]], [[0, 1, 2], [4, 6, 3], [7, 5, 8]], [[4, 1, 2], [0, 6, 3], [7, 5, 8]], [[4, 1, 2], [6, 0, 3], [7, 5, 8]], [[4, 1, 2], [6, 3, 0], [7, 5, 8]], [[4, 1, 0], [6, 3, 2], [7, 5, 8]], [[4, 0, 1], [6, 3, 2], [7, 5, 8]], [[4, 3, 1], [6, 0, 2], [7, 5, 8]], [[4, 3, 1], [6, 5, 2], [7, 0, 8]], [[4, 3, 1], [6, 5, 2], [0, 7, 8]], [[4, 3, 1], [0, 5, 2], [6, 7, 8]], [[0, 3, 1], [4, 5, 2], [6, 7, 8]], [[3, 0, 1], [4, 5, 2], [6, 7, 8]], [[3, 1, 0], [4, 5, 2], [6, 7, 8]], [[3, 1, 2], [4, 5, 0], [6, 7, 8]], [[3, 1, 2], [4, 0, 5], [6, 7, 8]], [[3, 1, 2], [0, 4, 5], [6, 7, 8]], [[0, 1, 2], [3, 4, 5], [6, 7, 8]]]
Error code: 0
```

<h3>
Functionality Tests
</h3>

<!-- grade_id: 8fbc24, points: 1, grade: true, solution: false -->
```python
## Test for state not correctly defined

incorrect_state = [[0,1,2],[2,3,4],[5,6,7]]
_,_,_,_,err = solvePuzzle(incorrect_state, lambda state: 0)
assert(err == -1)
```

<!-- grade_id: cell-5738a2124877e52e, points: 2, grade: true, solution: false -->
```python
## Heuristic function tests for misplaced tiles and manhattan distance

# Define the working initial states
working_initial_states_8_puzzle = ([[2,3,7],[1,8,0],[6,5,4]], [[7,0,8],[4,6,1],[5,3,2]], [[5,7,6],[2,4,3],[8,1,0]])

# Test the values returned by the heuristic functions
h_mt_vals = [7,8,7]
h_man_vals = [15,17,18]

for i in range(0,3):
    h_mt = heuristics[0](working_initial_states_8_puzzle[i])
    h_man = heuristics[1](working_initial_states_8_puzzle[i])
    assert(h_mt == h_mt_vals[i])
    assert(h_man == h_man_vals[i])

```

<!-- grade_id: cell-99c07d9541559b52, points: 4, grade: true, solution: false -->
```python
## A* Tests for 3 x 3 boards
## This test runs A* with both heuristics and ensures that the same optimal number of steps are found
## with each heuristic.

# Optimal path to the solution for the first 3 x 3 state
opt_path_soln = [[[2, 3, 7], [1, 8, 0], [6, 5, 4]], [[2, 3, 7], [1, 8, 4], [6, 5, 0]], 
                 [[2, 3, 7], [1, 8, 4], [6, 0, 5]], [[2, 3, 7], [1, 0, 4], [6, 8, 5]], 
                 [[2, 0, 7], [1, 3, 4], [6, 8, 5]], [[0, 2, 7], [1, 3, 4], [6, 8, 5]], 
                 [[1, 2, 7], [0, 3, 4], [6, 8, 5]], [[1, 2, 7], [3, 0, 4], [6, 8, 5]], 
                 [[1, 2, 7], [3, 4, 0], [6, 8, 5]], [[1, 2, 0], [3, 4, 7], [6, 8, 5]], 
                 [[1, 0, 2], [3, 4, 7], [6, 8, 5]], [[1, 4, 2], [3, 0, 7], [6, 8, 5]], 
                 [[1, 4, 2], [3, 7, 0], [6, 8, 5]], [[1, 4, 2], [3, 7, 5], [6, 8, 0]], 
                 [[1, 4, 2], [3, 7, 5], [6, 0, 8]], [[1, 4, 2], [3, 0, 5], [6, 7, 8]], 
                 [[1, 0, 2], [3, 4, 5], [6, 7, 8]], [[0, 1, 2], [3, 4, 5], [6, 7, 8]]]

astar_steps = [17, 25, 28]
for i in range(0,3):
    steps_mt, expansions_mt, _, opt_path_mt, _ = solvePuzzle(working_initial_states_8_puzzle[i], heuristics[0])
    steps_man, expansions_man, _, opt_path_man, _ = solvePuzzle(working_initial_states_8_puzzle[i], heuristics[1])
    # Test whether the number of optimal steps is correct and the same
    assert(steps_mt == steps_man == astar_steps[i])
    # Test whether or not the manhattan distance dominates the misplaced tiles heuristic in every case
    assert(expansions_man < expansions_mt)
    # For the first state, test that the optimal path is the same
    if i == 0:
        assert(opt_path_mt == opt_path_soln)

```

<!-- grade_id: cell-2981ba2cc7745c22, points: 8, grade: true, solution: false -->
```python
## A* Test for 4 x 4 board
## This test runs A* with both heuristics and ensures that the same optimal number of steps are found
## with each heuristic.

working_initial_state_15_puzzle = [[1,2,6,3],[0,9,5,7],[4,13,10,11],[8,12,14,15]]
steps_mt, expansions_mt, _, _, _ = solvePuzzle(working_initial_state_15_puzzle, heuristics[0])
steps_man, expansions_man, _, _, _ = solvePuzzle(working_initial_state_15_puzzle, heuristics[1])
# Test whether the number of optimal steps is correct and the same
assert(steps_mt == steps_man == 9)
# Test whether or not the manhattan distance dominates the misplaced tiles heuristic in every case
assert(expansions_mt >= expansions_man)
```

<!-- grade_id: cell-21cada9978e9a1bb, points: 16, grade: true, solution: false -->
```python
## Puzzle solvability test

unsolvable_initial_state = [[7,5,6],[2,4,3],[8,1,0]]
_,_,_,_,err = solvePuzzle(unsolvable_initial_state, lambda state: 0)
assert(err == -2)
```

### RESOURCES

> GeeksforGeeks. (2023, February 24). 8 puzzle problem using branch and bound. GeeksforGeeks. Retrieved March 2023, from https://www.geeksforgeeks.org/8-puzzle-problem-using-branch-and-bound/ 

> H, V. M. (2022, January 20). Implementation of a star search algorithm in Python. VTUPulse. Retrieved March 2023, from https://www.vtupulse.com/artificial-intelligence/implementation-of-a-star-search-algorithm-in-python/ 

> Pochmann, S. (2016, September 29). Calculating the Manhattan distance in the eight puzzle. Stack Overflow. Retrieved March 2023, from https://stackoverflow.com/questions/39759721/calculating-the-manhattan-distance-in-the-eight-puzzle 

> Shekhar, R. (2023, January 23). [template] A* Search. Minerva Forum.  Retrieved March 2023, from 
https://sle-collaboration.minervaproject.com/?minervaNotebookId=clcxqpd14012i0jzebthh5b5z&userId=10068&name=Edith+Ngundi&avatar=https%3A//s3.us-east-1.amazonaws.com/picasso.fixtures/Edith_Ngundi_10068_2022-08-30T01%3A48%3A53.155Z&noPresence=1&readOnly=0&isInstructor=0&signature=e89d6a6e927744650e081438ac3e7bc5418ea683f0d93284e5862af8e67de8d7